Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

```python
import os
import os.path
import io
import sys
from contextlib import redirect_stdout

datadir = "publicdata"
```

**Q** Write a function

    lastWords(filepath)

that creates and returns a list of strings consisting of the last word in each of the lines of the text file whose location is given by `filepath`.  The items in the list should not include any newline or whitespace as part of any of the words, and there should not be any items in the list for any lines consisting of all whitespace.  Make sure you remove any trailing punctuation.

<!-- grade_id: cell-05079d58ab7fbd80, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
filepath = os.path.join(datadir, "tennyson.txt")
print(lastWords(filepath))
```

<!-- grade_id: cell-ad7839c06c090b55, points: 3, grade: true, solution: false -->
```python
# Testing cell

assert lastWords(os.path.join(datadir, "tennyson.txt")) == ['bell', 'dark', 'farewell', 'embark']

```

**Q:** Write a function

    printType(path)

that prints the string "Is a directory" (without quotes) if the given `path` is a path to a directory, and prints "Is a file" if the given `path` is a regular file, and prints "Does not exist" if the path does not yield a traversal to an actual file or directory in the file system.  Note that the printed strings must match *exactly* in case and character sequence to the given specification.

<!-- grade_id: cell-27e2396f330b04d3, points: 0, grade: false, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
printType(".")
```

<!-- grade_id: cell-69dfc208049b8488, points: 3, grade: true, solution: false -->
```python
# Testing cell

path1 = "./other.ipynb"
path2 = os.path.abspath("publicdata")
path3 = "publicdata/foo.txt"

result = io.StringIO()
with redirect_stdout(result):
    printType(path1)
assert result.getvalue() == "Is a file\n"

result = io.StringIO()
with redirect_stdout(result):
    printType(path2)
assert result.getvalue() == "Is a directory\n"

result = io.StringIO()
with redirect_stdout(result):
    printType(path3)
assert result.getvalue() == "Does not exist\n"
```

**Q** Write a function

    vectorLengths(data)

that generates and returns a new list, where each element in the created list contains the integer length of the corresponding element from `data`.  So, for example, 

    vectorLengths(['this', 'is', 'a', 'short', 'list'])

yields the list `[4, 2, 1, 5, 4]`.

<!-- grade_id: cell-cfc3510a27b1b462, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Q** Write a function 

    apply_func(f,data) 

that modifies `data` by applying a function `f` to each item in `data`.  Note that this modifies `data` in place, and does not create a new list, nor return a value from the function.

<!-- grade_id: cell-b20194501fa27e7e, points: 2, grade: true, solution: true -->
```python
# Solution cell

# YOUR CODE HERE
raise NotImplementedError()
```

**Q** Write a general function

    apply_binary(vec1,vec2,f)

that applies a binary function `f` to two lists `vec1` and `vec2` (of the same length) and returns the resulting list.

<!-- grade_id: cell-157323bf54c7cabb, points: 2, grade: true, solution: true -->
```python
# Solution cell
# YOUR CODE HERE
raise NotImplementedError()
```